# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
YOUR_ID = "" # example: "rvuduc3"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

---

## Part 2: Quick review of Python data structures

In the [Python Bootcamp by Chris Simpkins](http://www.cc.gatech.edu/~simpkins/teaching/python-bootcamp/august2016.html), you learned about several basic [data structures for storing collections](http://www.cc.gatech.edu/~simpkins/teaching/python-bootcamp/slides/data-structures.html), including

- [lists](https://docs.python.org/3.5/library/stdtypes.html#lists),
- [dictionaries](https://docs.python.org/3.5/library/stdtypes.html#mapping-types-dict), and
- [sets](https://docs.python.org/3.5/library/stdtypes.html#set-types-set-frozenset).

In this part, let's build some intuition for the basic tradeoffs among these.

Let's start by considering the following program fragment.

```python
import random

def foo (n, k):
    return [random.randint (0, n) for _ in range (k)]

random.seed (0xDEADBEEF)
x = foo (1000000000, 1000000)

print (x[1:10])
```

> _(1 point)_ **Question 1.** Describe the contents of `x`. That is, what does `foo (n, k)` compute?

<!-- grade_id: what_does_foo_dom, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

> _(5 points)_ **Question 2.** Using only lists---that is, _without_ using dictionaries, sets, or other auxiliary data structures other than primitive variables or other lists---write a function that counts the number of unique entries of a given list.

<!-- grade_id: count_unique_lists, points: 2, grade: false, solution: true -->
```python
def count_unique__using_lists (x):
    """Returns the number of unique entries in `x`."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: check_count_unique_lists, points: 5, grade: true, solution: false -->
```python
n_unique = count_unique__using_lists (x)
print ('n =', len (x), 'and n_unique =', n_unique)
assert n_unique == 999492
print ("(Passed.)")
%timeit count_unique__using_lists (x)
```

> _(5 points)_ **Question 3.** Repeat the same exercise as above, but this time you may use dictionaries (`dict` in Python) in addition to lists.

<!-- grade_id: count_unique_dict, points: 2, grade: false, solution: true -->
```python
def count_unique__using_dicts (x):
    """Returns the number of unique entries in `x`."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_count_unique_dict, points: 5, grade: true, solution: false -->
```python
n_unique = count_unique__using_dicts (x)
print ('n =', len (x), 'and n_unique =', n_unique)
assert n_unique == 999492
print ("(Passed.)")
%timeit count_unique__using_dicts (x)
```

> _(5 points)_ **Question 4.** Repeat the same exercise one last time, but this time you should try to use a Python `set` and avoid using dictionaries.

<!-- grade_id: count_unique_set, points: 2, grade: false, solution: true -->
```python
def count_unique__using_sets (x):
    """Returns the number of unique entries in `x`."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: test_count_unique_set, points: 5, grade: true, solution: false -->
```python
n_unique = count_unique__using_sets (x)
print ('n =', len (x), 'and n_unique =', n_unique)
assert n_unique == 999492
print ("(Passed.)")
%timeit count_unique__using_sets (x)
```

> _(2 points)_ **Question 5.** You should have found that the three methods differ both in terms of their simplicity and in terms of their performance. Can you think of a scenario in which the results would have been different? (For instance, suppose you found above that a dictionary-based method outperformed a set-based method. Then describe an example problem in which you might expect a set-based method to instead outperform a dictionary-based method.)

<!-- grade_id: ds_analyze, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE