# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Processing unstructured text

A common task in data mining is processing unstructured text. This lab reviews some relevant material from the [Python bootcamp](http://www.cc.gatech.edu/~simpkins/teaching/python-bootcamp/august2016.html) in Part 1 and then, in Part 2, introduces a classic tool for more complex pattern matching known as the [regular expression](https://docs.python.org/3/howto/regex.html).

# Part 1: Simple string processing review

Do the following exercises.

**Exercise 1.** Write a function that returns `True` if a given input string is non-empty _and_ consists entirely of digits.

<!-- grade_id: is_all_digits, points: 0, grade: false, solution: true -->
```python
def is_all_digits (s):
    """Returns True only if the input string contains only digits."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: is_all_digits__test, points: 1, grade: true, solution: false -->
```python
assert is_all_digits ('22385')
assert not is_all_digits ('-238')
assert not is_all_digits ('')
assert not is_all_digits ('54f')
```

**Exercise 2.** Use your `is_all_digits()` function to create a new function that checks whether a given input string is a properly formatted social security number, i.e., has the pattern, `XXX-XX-XXXX`, _including_ the separator dashes, where each `X` is a digit.

<!-- grade_id: is_ssn, points: 0, grade: false, solution: true -->
```python
def is_ssn (s):
    """Returns True if the input string is a social security 
    number of the form, `XXX-XX-XXXX` (including dashes).
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: is_ssn__test1, points: 1, grade: true, solution: false -->
```python
assert is_ssn ('832-38-1847')
assert not is_ssn ('832381847')
assert not is_ssn ('8323-8-1847')
assert not is_ssn ('abc-de-ghij')
```

<!-- grade_id: is_ssn__test2, points: 1, grade: true, solution: false -->
```python
orig_is_all_digits = is_all_digits
del is_all_digits
try:
    is_ssn ('832-38-1847')
except NameError:
    pass
else:
    raise AssertionError ('`is_ssn` should use `is_all_digits`.')
finally:
    is_all_digits = orig_is_all_digits
```

```python

```