# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Mining the web, part 2: Web APIs (6 point)

We hope the preceding exercise was painful: it is rough downloading raw HTML and trying to extract information from it!

> Can you think of any other reasons why scraping websites for data in this way is not a good idea?

Luckily, many websites provide an application programming interface (API) for querying their data or otherwise accessing their services from your programs. For instance, Twitter provides a web API for gathering tweets, Flickr provides one for gathering image data, and Github for accessing information about repository histories.

These kinds of web APIs are much easier to use than, for instance, the preceding technique which scrapes raw web pages and then has to parse the resulting HTML. Moreover, there are more scalable in the sense that the web servers can transmit structured data in a less verbose form than raw HTML. In Homework 1, you will apply the techniques below, as well as others, to write some Python scripts to interact with the Yelp! web API.

As a starting example, here is some code to look at the activity on Github related to our course's IPython notebook repository.

```python
import requests

response = requests.get ('https://api.github.com/repos/rvuduc/cse6040fa16labs/events')

print ("==> .headers:", response.headers, "\n")
```

Note the `Content-Type` of the response:

```python
print (response.headers['Content-Type'])
```

The response is in JSON format, which is an open format for exchanging semi-structured data. (JSON stands for **J**ava**S**cript **O**bject **N**otation.) JSON is designed to be human-readable and machine-readable, and maps especially well in Python to nested dictionaries. Let's take a look.

> See also [this tutorial](http://www.w3schools.com/json/) for a JSON primer. JSON is among _the_ universal formats for sharing data on the web; see, for instance, https://www.sitepoint.com/10-example-json-files/.

<!-- grade_id: inspect_example1, points: 1, grade: true, solution: true -->
```python
# Inspect the results of the first cell before moving on

# YOUR CODE HERE
raise NotImplementedError()
```

**Exercise 1** (2 points). It should be self-evident that the JSON response above consists of a sequence of records, which we will refer to as _events_. Each event is associated with an _actor_. Write some code to extract a dictionary of all actors, where the key is the actor's login and the value is the actor's URL.

<!-- grade_id: extract_actors, points: 0, grade: false, solution: true -->
```python
def extract_actors (json_github_events):
    """Given JSON records for events in a GitHub repo,
    returns a dictionary of the actors and their URLs.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: extract_actors_test, points: 2, grade: true, solution: false -->
```python
actor_urls = extract_actors (response.json ())

assert len (actor_urls) == 8
for actor, url in actor_urls.items ():
    print ('{}: {}'.format (actor, url))
    assert url == "https://api.github.com/users/{}".format (actor)
```

**Exercise 2** (3 points). Write some code that goes to each actor's URL and determines their name. If an actor URL is invalid, that actor should not appear in the output.

<!-- grade_id: get_names, points: 0, grade: false, solution: true -->
```python
def lookup_names (actor_urls):
    """Given a dictionary of (actor, url) pairs, looks up the JSON at
    the URL and extracts the user's name (if any). Returns a new
    dictionary of (actor, name) pairs.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: get_names_test, points: 3, grade: true, solution: false -->
```python
actor_names = lookup_names (actor_urls)

assert len (actor_names) == 8
for actor, name in actor_names.items ():
    print ("{}: {}".format (actor, name))
    
assert actor_names['rvuduc'] == 'Rich Vuduc (personal account)'
```