# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Linear Algebra Review

The goal of this lab (homework) is to force you to refresh your linear algebra chops. Before starting it, please read these [linear algebra review notes](https://t-square.gatech.edu/access/content/group/gtc-3bd6-e221-5b9f-b047-31c7564358b7/Kuang-2014-linalg-notes.pdf), which were originally written by [Da Kuang](http://math.ucla.edu/~dakuang/), who taught CSE 6040 in Fall 2014.

## Part 1: Plug and chug [20 points]

Please do the following exercises *by hand*. You might verify your answers with a calculator or other numerical tool, but we want you to show your work. You may typeset a response below using the given Markdown cell and Jupyter's embedded LaTeX support.

> You may find this website useful for generating LaTeX for math symbols: http://detexify.kirelabs.org/classify.html

The point of these exercises is just to reinforce some of the definitions from the notes. There shouldn't be anything tricky here.

> If you don't wish to typeset your responses below, you may write them out by hand and scan and upload them as PDF files to T-Square.

Let

$x = \left(
       \begin{array}{c}
         1 \\
         2 \\
         3
       \end{array}
     \right)$
     
and

$A = \left(
       \begin{array}{ccc}
         2 & 1 & 0 \\
         4 & 3 & 1 \\
         6 & 7 & 8
       \end{array}
     \right)$.

Compute the following. (As an example, we've created a dummy "Exercise 0" that demonstrates a few useful typesetting constructs.)

**Exercise 0** (0 points). Compute $x + x$ of the numerical $x$ given above.

**Sample answer 1.**

$$
  x + x
  = 2x
  = 2 \cdot \left(\begin{array}{c} 1 \\ 2 \\ 3 \end{array}\right)
  = \left(\begin{array}{c} 2 \\ 4 \\ 6 \end{array}\right)
  .
$$

** Sample answer 2.**

$$
\begin{eqnarray}
  x + x & = & 2 \cdot x \\
        & = & 2 \cdot \left(\begin{array}{c} 1 \\ 2 \\ 3 \end{array}\right) \\
        & = & \left(\begin{array}{c} 2 \\ 4 \\ 6 \end{array}\right) .
\end{eqnarray}
$$

Now it's your turn!

**Exercise 1** (2 points). Compute the _inner-product_ (or _dot-product_), $x^T x$, of the numerical $x$ given above.

<!-- grade_id: dot_product, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 2** (2 points). Compute the _outer-product_, $x x^T$, of the numerical $x$ given above.

<!-- grade_id: outer_product, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 3** (2 points). Compute the _vector 1-norm_, $\| x \|_1$, of the numerical $x$ given above.

<!-- grade_id: vec_1norm, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 4** (2 points). Compute the _vector 2-norm_ (or _Euclidean length_ or _distance_), $\| x \|_2$, of the numerical $x$ given above.

<!-- grade_id: vec_2norm, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 5** (2 points). Compute the _infinity-norm_, $\| x \|_\infty$, of the numerical $x$ given above.

<!-- grade_id: vec_infnorm, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 6** (2 points). Compute the Frobenius norm, $\| A \|_F$, of the numerical matrix $A$ given above.

<!-- grade_id: frob_norm, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 7** (2 points). Compute the product, $A^T A$, of the numerical matrix $A$ given above.

<!-- grade_id: ata, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 8** (2 points). Compute the product, $A A^T$, of the numerical matrix $A$ given above.

<!-- grade_id: aat, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 9** (5 points). The _trace_ of a matrix $M$, or $\mathrm{trace}(M)$, is defined to be the sum of $M$'s diagonal entries, i.e., $\mathrm{trace}(M) = \sum_{i} m_{ii}$.

Let $A$ be any real-valued matrix of size $m \times n$. Show that $\mathrm{trace}(A^T A) = \mathrm{trace}(A A^T) = \| A \|_F^2$.

> In this question, show the desired claim for _any_ $m \times n$ matrix $A$, **not** for the numerical matrix $A$ given above.

<!-- grade_id: trace, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 10** (5 points). Suppose $x$ is now unknown and $b = (1, 1, 1)^T$. Determine a vector $x$ such that $Ax = b$, given the numerical $3 \times 3$ matrix $A$ above. You may use _any_ method, but show your work. (Methods might include basic high school algebra starting from the definitions of a matrix-vector product, or Gaussian elimination, or LU factorization.)

<!-- grade_id: gauss_elim, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 2: A proof [5 points]

**Exercise 11** (5 points). Let $Q \in \mathcal{R}^{m \times m}$ be an orthogonal matrix, and let $A \in \mathcal{R}^{m \times n}$ be some matrix. Prove that $\|QA\|_F^2 = \|A\|_F^2$, that is, left-multiplying by an orthogonal matrix does not change the Frobenius norm.

> Hint: Consider the definitions of an orthogonal matrix, Frobenius norm, vector 2-norm, and the results of Exercise 9.

<!-- grade_id: proof, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Part 3: Matrix-vector products in SQL [5 points]

Suppose you have a database with two tables, `A` and `X`, whose schema (equivalent `create` statements) are as follows.

```sql
  create table A (I integer, J integer, V real);
  create table X (I integer, V real);
```

Assume that the SQL statement,

```sql
  select min (I), max (I), min (J), max (J) from A;
```

returns the values,

| `min (I)` | `max (I)` | `min (J)` | `max (J)` |
|:---------:|:---------:|:---------:|:---------:|
|   $i_0$   |    $m$    |   $j_0$   |    $n$    |

where $i_0$, $m$, $j_0$, and $n$ are integers such that $1 \leq i_0 \leq m$ and $1 \leq j_0 \leq n$. Similarly, you may also assume that

```sql
  select min (I), max (I) from X;
```

returns the values:

| `min (I)` | `max (I)` |
|:---------:|:---------:|
|   $k_0$   |    $n$    |

where $1 \leq k_0 \leq n$.

Observe that the table `A` is like an $m \times n$ matrix $A = (a_{ij})$ and `X` is like a $n$-vector $x = (x_j)$, where particular values of `V` in `A` and `X` correspond to entries $a_{ij}$ of $A$ and $x_j$ of $x$. To complete the analogy between the database tables `(A, X)` and mathematical objects $(A, x)$, further assume that

- wherever $a_{ij} = 0$, the corresponding entries `(I, J)` of table `A` in the database is missing;
- wherever $x_i = 0$, the corresponding entry `I` of table `X` is also missing;
- wherever there are duplicate values of `(I, J)` in `A`, the true $a_{ij}$ is simply the sum of all the corresponding values of `V` in table `A`; and
- wherever there are duplicate values of `I` in `X`, the true $x_i$ is also simply the sum of all the corresponding values of `V` in table `X`.

**Exercise 12** (5 points). 

Write an SQL query (`select` statement) that computes the matrix-vector product, $A \cdot x$. Your statement should return a new table whose columns are named `I` and `V` and there are **no duplicates values in column `I` of the output**.

> You should write the SQL statement using Markdown in the following cell. That is, you need to write a valid `select` statement but you do not need to "execute" it.

<!-- grade_id: sql_matvec, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE