# Important note!

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
YOUR_ID = "" # Please enter your GT login, e.g., "rvuduc3" or "gtg911x"
COLLABORATORS = [] # list of strings of your collaborators' IDs
```

<!-- grade_id: who__test, points: 1, grade: true, solution: false -->
```python
import re

RE_CHECK_ID = re.compile (r'''[a-zA-Z]+\d+|[gG][tT][gG]\d+[a-zA-Z]''')
assert RE_CHECK_ID.match (YOUR_ID) is not None

collab_check = [RE_CHECK_ID.match (i) is not None for i in COLLABORATORS]
assert all (collab_check)

del collab_check
del RE_CHECK_ID
del re
```

**Jupyter / IPython version check.** The following code cell verifies that you are using the correct version of Jupyter/IPython.

```python
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

# Part 1: Intro to Numpy/Scipy [11 points]

[Numpy](http://www.numpy.org/) is a Python module that provides fast primitives for multidimensional arrays. It's well-suited to implementing numerical linear algebra algorithms, and for those can be much faster than Python's native list and dictionary types when you only need to store and operate on numerical data.

Some of the material from this lesson is lifted from the following more comprehensive tutorial: [link](http://www.scipy-lectures.org/intro/numpy/index.html)

**Quick demo.** The recommended importing idiom is:

```python
import numpy as np
```

Numpy provides some natural types and operations on arrays. For instance:

```python
a_1d = np.array ([0, 1, 2, 3]) # a vector
print (a_1d)

b_1d = np.array ([4, 5, 6, 7]) # another vector
print (b_1d)

print (a_1d + b_1d)
```

```python
print (5*a_1d)
```

```python
print (a_1d**2)
```

**Getting help.** By the way, if you need help getting documentation from within this notebook, here are some handy shortcuts.

```python
# Append '?' to get help on a specific routine
np.array?
```

```python
# Search for key text
np.lookfor ("creating array")
```

```python
# Wildcard search
np.con*?
```

## Why bother with Numpy? A motivating example

We already have lists and dictionary types, which are pretty easy to use and very flexible. So why bother with this special type?

One reason to consider Numpy is that it "can be much faster," as noted above. But how much faster is that?

```python
n = 1000000
```

```python
L = range (n)
%timeit [i**2 for i in L]
```

```python
np.arange (10) # Moral equivalent to `range`
```

```python
A = np.arange (n)
%timeit A**2
```

**Exercise 1** (1 point). Recall the definition of the _2-norm_ of a vector (or _Euclidean length_ of a vector) from [Da Kuang's notes on linear algebra](https://t-square.gatech.edu/access/content/group/gtc-3bd6-e221-5b9f-b047-31c7564358b7/Kuang-2014-linalg-notes.pdf). Compare its speed when using native Python lists versus Numpy arrays.

> Hint: For Numpy, look for a routine that computes the norm.

<!-- grade_id: norm, points: 1, grade: true, solution: true -->
```python
from random import gauss # Generates random numbers from a Gaussian
from math import sqrt # Computes the square root of a number

n = 1000000
X_py = [gauss (0, 1) for i in range (n)]
X_np = np.array (X_py)

print ("==> Native Python lists:")
%timeit sqrt (sum ([x**2 for x in X_py]))

print ("\n==> Numpy:")
# YOUR CODE HERE
raise NotImplementedError()
```

## Creating multidimensional arrays

Beyond simple arrays, Numpy supports multidimensional arrays. To do more than one dimension, call `numpy.array()` but nest each new dimension within a list.

Huh? Let's look at some examples.

```python
# Create a two-dimensional array of size 3 rows x 4 columns:
B = np.array([[0, 1, 2, 3],
              [4, 5, 6, 7],
              [8, 9, 10, 11]])

print (B)
print (B.ndim) # What does this do?
print (B.shape) # What does this do?
print (len (B)) # What does this do?
```

```python
C1 = [[0, 1, 2, 3],
      [4, 5, 6, 7],
      [8, 9, 10, 11]]

C2 = [[12, 13, 14, 15],
      [16, 17, 18, 19],
      [20, 21, 22, 23]]

C = np.array ([C1, C2])

print (C)
print (C.ndim)
print (C.shape)
print (len (C))
```

Besides `arange()`, you can also define an interval and a number of points. What does the following code do?

```python
print (np.linspace (0, 1, 10))
```

```python
print (np.linspace (0, 1, 10, endpoint=False))
```

There are routines for creating various kinds of structured matrices as well, which are similar to those found in [MATLAB](http://www.mathworks.com/products/matlab/) and [Octave](https://www.gnu.org/software/octave/).

```python
print (np.ones ((3, 4)))
```

```python
print (np.zeros ((3, 4)))
```

```python
print (np.eye (3))
```

```python
print (np.diag ([1, 2, 3]))
```

**Exercise 2** (1 point). The following code creates an identity matrix in two different ways, which are found to be equal according to the assertion. But in fact there is a subtle difference between the `I` and `I_u` matrices created below; can you spot it? Explain that difference in the `YOUR ANSWER HERE` (Markdown) cell that follows the code.

```python
n = 3
I = np.eye (n)

print ("==> I = eye(n):")
print (I)

u = [1] * n
I_u = np.diag (u)

print ("\n==> u:\n", u)
print ("==> I_u = diag (u):\n", I_u)

assert np.all (I_u == I)
```

<!-- grade_id: explain_diff, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

You can also create empty (uninitialized) arrays. What does the following produce?

```python
A = np.empty ((3, 4)) # An "empty" 3 x 4 matrix
print (A)
```

## Indexing and slicing

The usual 0-based slicing and indexing notation you know and love from lists is also supported for Numpy arrays. In the multidimensional case,  including their natural multidimensional analogues with index ranges separated by commas.

```python
# Recall: C
print (C)
```

What part of C will the following slice extract? Run the code to find out.

```python
print (C[0, 2, :])
```

What will the following slice return? Run the code to find out.

```python
print (C[1, 0, ::-1])
```

**Exercise 3** (5 points). Consider the following $6 \times 6$ matrix, which has 4 different subsets highlighted.

<img src="slicing-exercise.png" alt="Exercise: Extract these slices" width="240">

Write some code to generate this matrix, named `Z`. Then, for each subset illustrated above, write an indexing or slicing expression that extracts the subset. Store the result of each slice into `Z_green`, `Z_red`, `Z_orange`, and `Z_cyan`.

<!-- grade_id: create_Z, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: check_Z, points: 5, grade: true, solution: false -->
```python
print ("==> Z:\n", Z)
assert (Z == np.array ([np.arange (0, 6),
                        np.arange (10, 16),
                        np.arange (20, 26),
                        np.arange (30, 36),
                        np.arange (40, 46),
                        np.arange (50, 56)])).all ()

print ("\n==> Orange slice:\n", Z_orange)
assert (Z_orange == np.array ([3, 4])).all ()

print ("\n==> Red slice:\n", Z_red)
assert (Z_red == np.array ([2, 12, 22, 32, 42, 52])).all ()

print ("\n==> Cyan slice:\n", Z_cyan)
assert (Z_cyan == np.array ([[44, 45], [54, 55]])).all ()

print ("\n==> Green slice:\n", Z_green)
assert (Z_green == np.array ([[20, 22, 24], [40, 42, 44]])).all ()
```

## Slices are views

To help save memory, when you slice a Numpy array, you are actually creating a _view_ into that array. That means modifications through the view will modify the original array.

```python
print ("==> Recall C: %s" % str (C.shape))
print (C)
```

```python
C_view = C[1, 0::2, 1::2] # Question: What does this produce?
print ("==> C_view: %s" % str (C_view.shape))
print (C_view)
```

```python
C_view[:, :] = -C_view[::-1, ::-1] # Question: What does this do?
print (C_view)
```

```python
print (C)
```

You can force a copy using the `.copy()` method:

```python
C_copy = C[1, 0::2, 1::2].copy ()
C_copy[:, :] = -C_copy[::-1, ::-1]

print ("==> C_view:")
print (C_view)

print ("\n==> C_copy:")
print (C_copy)
```

And to check whether two Numpy array variables point to the same object, you can use the `numpy.may_share_memory()` function:

```python
print ("C and C_view share memory: %s" % np.may_share_memory (C, C_view))
print ("C and C_copy share memory: %s" % np.may_share_memory (C, C_copy))
```

**Exercise 4** (3 points). Complete the prime number sieve algorithm, which is illustrated below.

<img src="prime-sieve.png" alt="Exercise: Extract these slices" width="480">

That is, given a positive integer $n$, the algorithm iterates from $i \in \{2, 3, 4, \ldots, \left\lfloor\sqrt{n}\right\rfloor\}$, repeatedly "crossing out" values that are strict multiples of $i$. "Crossing out" means maintaining an array of, say, booleans, and setting values that are multiples of $i$ to `False`.

<!-- grade_id: prime_sieve, points: 0, grade: false, solution: true -->
```python
from math import sqrt

def sieve (n):
    """
    Returns the prime number 'sieve' shown above.
    
    That is, this function returns an array `X[0:n+1]`
    such that `X[i]` is true if and only if `i` is prime.
    """
    is_prime = np.empty (n+1, dtype=bool) # the "sieve"

    # Initial values
    is_prime[0:2] = False # {0, 1} are _not_ considered prime
    is_prime[2:] = True # All other values might be prime

    # Implement the sieving loop
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return is_prime

# Prints your primes
print ("==> Primes through 20:\n", np.nonzero (sieve (20))[0])
```

<!-- grade_id: prime_sieve_test, points: 3, grade: true, solution: false -->
```python
is_prime = sieve (20)
assert len (is_prime) == 21
assert (is_prime == np.array ([False, False, True, True, False, True, False, True, False, False, False, True, False, True, False, False, False, True, False, True, False])).all ()
```

## Indirect addressing

Two other common ways to index a Numpy array are to use a boolean mask or to use a set of integer indices.

```python
np.random.seed(3)
x = np.random.randint(0, 20, 15) # 15 random ints in [0, 20)
print (x)
```

Before looking at how to use a boolean mask for indexing, let's create one.

**Exercise 5** (1 point). Given the input array, `x[:]`, above, create an array, `mask_mult_3[:]` such that `mask_mult_3[i]` is true only if `x[i]` is a positive multiple of 3.

<!-- grade_id: mask_mult_3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: mask_mult_3_test, points: 1, grade: true, solution: false -->
```python
print ("x:", x)
print ("mask_mult_3:", mask_mult_3)
print ("==> x[mask_mult_3]:", x[mask_mult_3])

inv_mask_mult_3 = np.invert (mask_mult_3)
assert ((x[mask_mult_3] % 3) == np.zeros (sum (mask_mult_3))).all ()
assert (((x[inv_mask_mult_3] % 3) != np.zeros (sum (inv_mask_mult_3))) | (x[inv_mask_mult_3] == 0)).all ()
```

```python
# Pull out an arbitrary subset of elements
inds = np.array ([3, 7, 8, 12])
print (x[inds])
```